# Tutorial: Lambda-Funktionen

In diesem Tutorial werden Sie die grundlegende Funktionsweise von Lambda-Funktionen kennenlernen. Dies beinhaltet das Erstellen und das Aufrufen von Lambda-Funktionen.

## Einführung

Lambda-Funktionen sind anonyme Funktionen in Python. Sie sind nützlich für kurze, einfache Operationen. Eine Lambda-Funktion wird mit dem Schlüsselwort `lambda` definiert. Im folgenden Beispiel wird eine Lambda-Funktion erstellt, die 2 Zahlen addiert.

```python
lambda x, y: x + y
```

Jetzt wird die Lambda-Funktion mit 2 Zahlen aufgerufen:

```python
print(f'4 + 8 = {(lambda x, y: x + y)(4, 8)}')
```

```python
lambda bound_var1, bound_var2, ...: body
```

- `lambda`: Keyword, mit dem eine Lambda Funktionsdefinition eingeleitet wird.
- `bound_var`: Variablen, die nur innerhalb der Lambda Funktion gelten. Diese Variablen sind dann auch nur dort verwendbar.
- `body`: Enthält *einen* Ausdruck, der den Rückgabewert dieser Funktion bestimmt.

Typischerweise werden Lambda-Funktionen direkt verwendet oder an andere Funktionen übergeben. Sie könnten zwar auch auf Variablen zugewiesen und später verwendet werden, dann ist allerdings eine "klassische" Funktionsdefinition sinnvoller.

## Verwendung von Lambda in Funktionen

Lambda-Funktionen können in Funktionen wie `map()`, `filter()` und `sorted()` verwendet werden.
Hier ein kleines Beispiel, in welchem die Quadratzahlen von verschiedenen Zahlen berechnet werden.

```python
list(map(lambda x: x**2, [1, 2, 3, 4]))
```

Hier ein kleines Beispiel, in welchem alle Wörter mit einem `"A"` am Anfang herausgefiltert werden.

```python
fruits = ["Apple", "Banana", "Orange", "Pear", "Apricot"]
list(filter(lambda s: s.startswith("A"), fruits))
```

Sie können Lambda-Funktionen in benutzerdefinierten Funktionen integrieren.
Im folgenden Beispiel wird eine Funktion erstellt und mit der gleichen Lambda-Funktion wie oben aufgerufen.

```python
def apply_operation(x, operation): 
    return operation(x)

apply_operation(5, lambda x: x**2)
```

## Ihre Aufgabe

Schreiben Sie eine Funktion namens `filter_words_by_length`, die eine Liste von Wörtern und eine minimale Längenbeschränkung als Parameter erhält. Die Funktion sollte eine neue Liste zurückgeben, die nur die Wörter enthält, deren Länge größer oder gleich der angegebenen Längenbeschränkung ist. Verwenden Sie eine Lambda-Funktion, um diese Aufgabe zu lösen.

Beispiel:
```python
def filter_words_by_length(word_list, min_length):
    return list(filter(Lambda_Funktion))

words = ["apple", "banana", "kiwi", "orange", "grape", "fun"]
min_length = 5
result = filter_words_by_length(words, min_length)
# result = ['apple', 'banana', 'orange', 'grape']
```

```python
words = ["apple", "banana", "kiwi", "orange", "grape", "fun"]
min_length = 5
```

<!-- grade_id: cell-8adec7b63b20f201, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.
```

<!-- grade_id: cell-5e964c7e24aff6a1, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(filter_words_by_length)

args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']

assert pytest.main([*args, pathlib.Path('tests') / 'Lambda_functions' / 'filter_words.yml']) is pytest.ExitCode.OK
```