# Slicing


### Aufgabe 1 - Palindrom Erkennung

Ein Palindrom ist ein Wort, das vorwärts und rückwärts gelesen dasselbe ist, wie zum Beispiel **"otto"** oder **"rentner"**. In dieser Aufgabe werden wir Slicing in Python anwenden, um zu überprüfen, ob ein gegebenes Wort ein Palindrom ist.


Sie erhalten ein `string` auf der Variablen `moegliches_palindrom` und es ist Ihre Aufgabe zu bestimmen, ob es sich bei dem String um ein Palindrom handelt:

- schreiben Sie ihr Ergebnis, ob es sich um ein Palindrom handelt oder nicht, als `bool` auf die Variable `ist_palindrom`


```python
moegliches_palindrom = "elle"
```

### Aufgabe 2

7 Boom ist ein Spiel, bei dem die Spieler abwechselnd natürliche Zahlen aufzählen müssen. Dabei wird anstelle jeder Zahl, die durch 7 teilbar ist oder die die Ziffer 7 enthält, "Boom" gesagt.


Sie erhalten eine Liste der natürliche Zahlen $\mathbb{N}_{0}$, von 0 bis 100 in der Variablen `boom_7`:

- greifen Sie durch Slicing Operationen auf die Liste `boom_7` zu, und verändern Sie die Liste so, dass sie den Regeln von 7 Boom folgt. (e.g `[5,6,7]` $\to$ `[5,6,'Boom']`)


```python
boom_7 = [ 0,  1,  2,  3,  4,  5,  6,  7,  8,  9, 10, 11, 12, 13, 14, 15, 16,
       17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33,
       34, 35, 36, 37, 38, 39, 40, 41, 42, 43, 44, 45, 46, 47, 48, 49, 50,
       51, 52, 53, 54, 55, 56, 57, 58, 59, 60, 61, 62, 63, 64, 65, 66, 67,
       68, 69, 70, 71, 72, 73, 74, 75, 76, 77, 78, 79, 80, 81, 82, 83, 84,
       85, 86, 87, 88, 89, 90, 91, 92, 93, 94, 95, 96, 97, 98, 99, 100]
```

### Aufgabe 3

Das Zusammenführen von gleichlangen Listen kann nach zwei Verfahren durchgeführt werden.

- Das _'Block-Verfahren'_, bei dem eine Liste nach der anderen zu einer neuen Liste kombiniert werden (e.g `[1,2,3]` und `[a,b,c]` $\to$ `[1,2,3,a,b,c]`).
- Das _'Reißverschluss-Verfahren'_, bei dem jeweils abwechselnd Elemente der beiden Listen zur neuen Liste hinzugefügt werden (e.g `[1,2,3]` und `[a,b,c]` $\to$ `[1,a,2,b,3,c]`).


Sie erhalten zwei gleichlange Listen `liste_1` und `liste_2`, sowie eine Liste von gleichlangen Listen `liste_von_listen`:


```python
liste_1 = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j', 'k', 'l', 'm', 'n', 'o', 'p', 'q', 'r', 's', 't', 'u', 'v', 
           'w', 'x', 'y', 'z']
liste_2 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26]

liste_von_listen = [liste_1, liste_2, 
                    [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121, 144, 169, 
                     'N', 'O', 'P', 'Q', 'R', 'S', 'T', 'U', 'V', 'W', 'X', 'Y', 'Z']]
```

##### Aufgabenstellung

- implementieren Sie das _'Reißverschluss-Verfahren'_ unter Anwendung von Slicing um `liste_1` und `liste_2` zu einer neuen Liste `beispiel_verbundene_liste` zu kombinieren.
- implementieren Sie eine Verallgemeinerung, die unter Anwendung von Slicing all Listen in `liste_von_listen` zu einer neuen Liste `verbundene_liste` kombiniert. Sie können hierbei voraussetzen, dass alle Listen in `liste_von_listen` die gleiche Länge haben.

_Hinweis_: Sie können auch nur die Verallgemeinerte Version implementieren und auf `liste_1` und `liste_2` anwenden, aber es empfiehlt sich erst mit der spezifischen Version zu beginnen, und sie von da ausgehend zu verallgemeinern. Beachten Sie jedoch, dass ihre Zelle zur Erfüllung der Tests beide Variablen ausgeben muss.


<!-- grade_id: cell-eeec10dd26466871, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.
# Aufgabe 1 – Palindrom-Erkennung
ist_palindrom = moegliches_palindrom == moegliches_palindrom[::-1]

# Aufgabe 2 – 7 Boom (Slicing + Listenänderung)
boom_7 = list(range(101))
for i in range(len(boom_7)):
    if boom_7[i] % 7 == 0 or '7' in str(boom_7[i]):
        boom_7[i] = 'Boom'

# Aufgabe 3 – Reißverschluss-Verfahren (Slicing-Variante)
liste_1 = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j', 'k', 'l', 'm', 'n', 'o', 'p', 'q', 'r', 's', 't', 'u', 'v', 
           'w', 'x', 'y', 'z']
liste_2 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26]

# Beispielhafte Reißverschluss-Zusammenführung mit Slicing
beispiel_verbundene_liste = [None] * (len(liste_1) + len(liste_2))
beispiel_verbundene_liste[::2] = liste_1
beispiel_verbundene_liste[1::2] = liste_2

# Verallgemeinerte Version für liste_von_listen mit Slicing
liste_von_listen = [liste_1, liste_2, 
                    [1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121, 144, 169, 
                     'N', 'O', 'P', 'Q', 'R', 'S', 'T', 'U', 'V', 'W', 'X', 'Y', 'Z']]

verbundene_liste = []
for i in range(len(liste_von_listen[0])):
    for lst in liste_von_listen:
        verbundene_liste.append(lst[i])

```

<!-- grade_id: cell-d23c35a7e293bb99, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(_i)
cases = pathlib.Path('tests') / 'Slicing'
args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']

assert all(
    pytest.main([*args, cases / f'{subtask}.yml']) is pytest.ExitCode.OK
    for subtask in ['ist_palindrom', 'boom_7', 'verbundene_listen_beispiel', 'verbundene_listen']
)
```

Output:
```
The following submission will be tested:

# Bitte schreiben Sie hier Ihren Programmcode.
# Aufgabe 1 – Palindrom-Erkennung
ist_palindrom = moegliches_palindrom == moegliches_palindrom[::-1]

# Aufgabe 2 – 7 Boom (Slicing + Listenänderung)
boom_7 = list(range(101))
for i in range(len(boom_7)):
    if boom_7[i] % 7 == 0 or '7' in str(boom_7[i]):
        boom_7[i] = 'Boom'

# Aufgabe 3 – Reißverschluss-Verfahren (Slicing-Variante)
liste_1 = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j', 'k', 'l', 'm', 'n', 'o', 'p', 'q', 'r', 's', 't', 'u', 'v', 
           'w', 'x', 'y', 'z']
liste_2 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26]

# Beispielhafte Reißverschluss-Zusammenführung mit Slicing
beispiel_verbundene_liste = [None] * (len(liste_1) + len(liste_2))
beispiel_verbundene_liste[::2] = liste_1
beispiel_verbundene_liste[1::2] = liste_2

# Verallgemeinerte Version für liste_von_listen mit Slicing
liste_von_listen = [liste_1, liste_2, 
```

```python

```