# Transponieren

In dieser Aufgabe schreiben Sie eine einfache Funktion zur Arbeit mit Matrizen. Sie kann elegant mit eingebauten Funktionen und indefiniten Argumenten gelöst werden.

## Grundlage

Die Transponierte einer Matrix $\mathbf{X}$ mit den Einträgen $X_{ij}$ ist die Matrix $\mathbf{X}^T$ mit den Einträgen $X_{ji}$, d.h. Spalten und Zeilen wurden vertauscht. Zum Beispiel:

$$\mathbf{X} = \begin{pmatrix}1 & 2 & 3 \\ 4 & 5 & 6 \end{pmatrix}; ~~~~~~ \mathbf{X}^T = \begin{pmatrix}1 & 4\\ 2 & 5 \\ 3 & 6 \end{pmatrix}$$

In Standard-Python lassen sich Matrizen als `tuple` von `tuple`s von numerischen Datentypen darstellen. Für das obige Beispiel wäre also:
```python
x = ((1, 2, 3), (4, 5, 6))
x_t = ((1, 4), (2, 5), (3, 6))
```

(Bemerkung: Wir implementieren das Transponieren mit Python-`builtins` rein zu Übungszwecken. Mit `numpy` als `np.Array`s dargestellte Matrizen können mit der `np.transpose`-Methode transponiert werden. Nutzen Sie jedoch für diese Aufgabe keine Pakete.)

## Ihre Aufgabe

Implementieren Sie eine Funktion `transpose`, die für eine als Tuple von Tuples dargestellte Matrix eine neue gleicherart dargestellte Matrix erzeugt, die die Transponierte des Arguments ist. Die Funktion sollte also folgende Signatur haben:
```python
Matrix = tuple[tuple[int, ...], ...]
transpose(matrix: Matrix) -> Matrix
```

Sowohl quadratische Matrizen mit $n$ Zeilen und $n$ Spalten als auch rechteckige Matrizen sollten sich transponieren lassen.

<!-- grade_id: cell-d4880b09e54c079f, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.
```

<!-- grade_id: cell-2e27054884c54b62, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(transpose)
cases = pathlib.Path('tests') / 'Transponieren'
args = ['-qq', '-x',  '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']

assert pytest.main([*args, cases / 'square_matrices.yml']) is pytest.ExitCode.OK
assert pytest.main([*args, cases / 'rectangular_matrices.yml']) is pytest.ExitCode.OK
```