# Multimap umdrehen

In dieser Aufgabe iterieren Sie über Schlüssel-Wert-Paare eines Dictionaries und erstellen ein neues Dictionary. Die Aufgabe ist schnell gestellt und gibt keinen Lösungsweg vor.

## Problemstellung

Dictionaries sind sogenannte *Mappings*. Dabei ist jedem Schlüssel genau ein Wert zugeordnet. Manchmal möchte man einem Schlüssel mehr als einen Wert zuordnen. Beispielsweise kann eine Person mehrere Module besuchen, an jedem Modul nehmen andererseits beliebig viele Personen teil. In Python könnte man dies repräsentieren, indem man jedem Schlüssel nicht nur einen Wert, sondern eine *Menge* (ein `set`) von Werten zuordnet.

Da sowohl Schlüssel eines Dictionaries als auch Elemente einer Menge unveränderlich ("immutable") sind, lässt sich diese Zuordnung stets umkehren. Beispielsweise haben Sie ein Dictionary, das Personen als Schlüssel und Mengen von besuchten Veranstaltungen als Werte enthält. Dann könnten Sie ein neues Dictionary erstellen, das die einzelnen Veranstaltungen als Schlüssel und die Mengen der jeweils teilnehmenden Personen als Werte enthält.

## Ihre Aufgabe

Sie erhalten ein Dictionary `multimap`. Darin sind Schlüsseln (als Strings) ein oder mehrere Werte (als Sets von Integern) zugeordnet.

Erstellen Sie ein Dictionary `reverse_multimap` mit der umgekehrten Zuordnung. Dabei sollen die Integer in den Mengen des ursprünglichen Dictionaries zu Schlüsseln im neuen Dictionary werden und die ursprünglichen Schlüssel als Elemente von Mengen gespeichert werden, die den jeweiligen Integern als Werte zugeordnet sind. 

Beispiel:
```python
multimap = {
    'a': {1, 2, 3},
    'b': {1, 3},
    'c': {2}
}

reverse_multimap = {
    1: {'a', 'b'},
    2: {'a', 'c'},
    3: {'a', 'b'}
}
```

Auch wenn jede Antwort, die die Tests besteht, die Aufgabe "zutreffend bearbeitet" und nicht weiter bewertet wird: Geben Sie sich Mühe, die klarste und idiomatischste Antwort zu finden.

```python
# variable initialization.
multimap = {
    'a': {1, 2, 3},
    'b': {1, 3},
    'c': {2}
}
```

<!-- grade_id: cell-2a4f82a0fcc67409, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.

reverse_multimap = {}

for key, values in multimap.items():
    for value in values:
        if value not in reverse_multimap:
            reverse_multimap[value] = set()
        reverse_multimap[value].add(key)


```

<!-- grade_id: cell-b9e264a484cf89bd, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(_i)

cases = pathlib.Path('tests') / 'Multimap_umdrehen'
args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']

assert pytest.main([*args, cases / 'reverse_multimap.yml']) is pytest.ExitCode.OK
```

Output:
```
The following submission will be tested:

# Bitte schreiben Sie hier Ihren Programmcode.

reverse_multimap = {}

for key, values in multimap.items():
    for value in values:
        if value not in reverse_multimap:
            reverse_multimap[value] = set()
        reverse_multimap[value].add(key)
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m                                    [100%][0m
```

```python

```