# Matrizenaddition

In dieser Aufgabe iterieren Sie über einfach verschachtelte Listen. Die Aufgabe ist schnell gestellt und lädt Sie dazu ein, sich eigene Gedanken zu machen, wie sie erweitert werden könnte.

## Problemstellung

Matrizen sind rechteckige Zahlenschemata einer gegebenen Form. In Python könnten Matrizen als Listen-von-Listen repräsentiert werden. Dabei repräsentiert jede innere Liste eine neue Zeile der Matrix. Hier ein Vergleich, wie man Matrizen in der Mathematik und in Python darstellen kann:

$$
  \begin{pmatrix} 1 & 2 & 3\\ 4 & 5 & 6 \\ 7 & 8 & 9\end{pmatrix}~~~~~~~~~~~~~~~~ =  ~~~~~~~~~~~~~~~~
  \begin{matrix}
  \texttt{[} & \\
  & \texttt{[1, 2, 3],} \\ & \texttt{[4, 5, 6],} \\ & \texttt{[7, 8, 9],} \\ 
  \texttt{]} &
  \end{matrix}
  $$

Matrizen erlauben jede Menge Rechenoperationen. Insbesondere können zwei Matrizen addiert werden, wenn sie die gleiche Form besitzen. Dabei werden die einzelnen Elemente addiert. Beispielsweise:

$$\begin{pmatrix} 1 & 2 & 3\\ 4 & 5 & 6\end{pmatrix} + \begin{pmatrix} 2 & 3 & 5\\ 7 & 11 & 13\end{pmatrix} = \begin{pmatrix} 3 & 5 & 8\\ 11 & 16 & 19\end{pmatrix}$$

Diesen Vorgang kann man auch in Python umsetzen.

## Ihre Aufgabe

Gegeben sind zwei Variablen: `matrix1` und `matrix2`. Beides sind Listen von Listen von Zahlen. Sie können davon ausgehen, dass die beiden Matrizen die gleiche Form haben, also die gleiche Anzahl an Zeilen und Spalten.
Schreiben Sie in eine Variable `summe_matrix` die Summe der beiden Matrizen. Dies ist — nach obigem Schema — eine Liste-von-Listen der gleichen Form, deren Elemente die Summe der entsprechenden Elemente der beiden Matrizen sind.

Auch wenn jede Antwort, die die Tests besteht, die Aufgabe "zutreffend bearbeitet" und nicht weiter bewertet wird: Geben Sie sich Mühe, die klarste und idiomatischste Antwort zu finden. Achten Sie insbesondere darauf, Ihre Schleifen möglichst frei von Zählindizes zu halten. Eingebaute Funktionen wie `zip` und `sum` können nützlich sein.

Ebenfalls unbewertet ist folgender Teil: Überlegen Sie, welche Probleme auftauchen könnten. Welche Annahmen haben Sie getroffen? Was passiert, wenn diese nicht erfüllt sind? Wie könnten Sie damit umgehen? Und zuletzt: Wie könnte man effizient eine beliebige Zahl von Matrizen addieren?

```python
# variable initialization.
matrix1 = [[1, 2, 3], [4, 5, 6], [7, 8, 9]]
matrix2 = [[3, 1, 2], [6, 4, 5], [9, 7, 8]]
```

<!-- grade_id: cell-fffc2826f313b01a, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.

# Berechne die Summe der beiden Matrizen
summe_matrix = [
    [a + b for a, b in zip(zeile1, zeile2)]
    for zeile1, zeile2 in zip(matrix1, matrix2)
]


```

<!-- grade_id: cell-c286e66e32ca6edb, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(_i)
cases = pathlib.Path('tests') / 'Matrizenaddition'
args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']

assert pytest.main([*args, cases / 'summe_matrix.yml']) is pytest.ExitCode.OK
```

Output:
```
The following submission will be tested:

# Bitte schreiben Sie hier Ihren Programmcode.

# Berechne die Summe der beiden Matrizen
summe_matrix = [
    [a + b for a, b in zip(zeile1, zeile2)]
    for zeile1, zeile2 in zip(matrix1, matrix2)
]
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m                      [100%][0m
```

```python

```