# Buchstabenmagie

In dieser Aufgabe dürfen Sie mit `for`-Schleifen Wörter verarbeiten, die aus einem gegebenen Alphabet bestehen. Dabei müssen Sie die in natürlicher Sprache gestellte Aufgabe in Programmcode "übersetzen".

## Problemstellung

Alice hat für ihre Geheimsprache ein eigenes Alphabet erstellt. Dieses Alphabet verwendet nur ausgewählte Buchstaben und diese in einer neuen Reihenfolge. Außerdem gibt Sie Ihnen ein in dieser Geheimsprache geschriebenes Wort. Sie beginnen mit der Kryptoanalyse und stellen sich folgende Fragen: Wie wurden die Buchstaben bei Erstellung des Alphabets vermischt? Welche Buchstaben des Alphabets enthält das Wort wie oft?

## Ihre Aufgabe

Sie erhalten zwei Variablen mit Strings namens `alphabet` und `wort`.
* `alphabet` enthält einige Buchstaben, jeden davon genau einmal. Allerdings sind die Buchstaben nicht unbedingt auf herkömmliche Weise ("lexikographisch") geordnet.
* `wort` ist ein String mit beliebigem Inhalt. Insbesondere sind ein paar Buchstaben des Alphabets enthalten, in beliebiger Reihenfolge und beliebig oft.

### a) Alphabet entwirren

Vergleichen Sie die Buchstaben im `alphabet` mit ihrer herkömmlichen Sortierung. Erstellen Sie eine Liste `indizes`, welche die Indizes enthält, an welchen die Buchstaben im Alphabet vorkommen würden, wenn dieses sortiert wäre. Beispielsweise:
```python
alphabet = 'CBDA'
indizes = [2, 1, 3, 0]
```

### b) Häufigkeit zählen

Nun analysieren Sie `wort`. Sie möchten zählen, wie häufig die Buchstaben aus `alphabet` in `wort` enthalten sind. Erstellen Sie dazu ein Dictionary, das als Schlüssel die Buchstaben in `alphabet` verwendet und als zugehörige Werte die Häufigkeit dieses Buchstabens in `wort`. Beispielsweise:
```python
alphabet = 'ABCDR'
wort = "ABRACADABRA" 

zeichen_anzahl = {'A': 5, 'B': 2, 'C': 1, 'D': 1, 'R': 2}
```

```python
# variable initialization.
alphabet = 'RADBC'
wort = 'ABRACADABRA'
```

<!-- grade_id: cell-e3ebc806e200e742, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.

# Teil a) Alphabet entwirren: Position jedes Buchstabens aus `alphabet` im sortierten Alphabet
sortiertes_alphabet = sorted(alphabet)
indizes = [sortiertes_alphabet.index(buchstabe) for buchstabe in alphabet]

# Teil b) Häufigkeit zählen
zeichen_anzahl = {buchstabe: wort.count(buchstabe) for buchstabe in alphabet}

sortiertes_alphabet = sorted(alphabet)
indizes = [sortiertes_alphabet.index(buchstabe) for buchstabe in alphabet]

zeichen_anzahl = {buchstabe: wort.count(buchstabe) for buchstabe in alphabet}

```

<!-- grade_id: cell-c92889a159aed308, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(_i)
cases = pathlib.Path('tests') / 'Buchstabenmagie'
args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']

assert pytest.main([*args, cases / 'indizes.yml']) is pytest.ExitCode.OK
assert pytest.main([*args, cases / 'zeichen_anzahl.yml']) is pytest.ExitCode.OK
```

Output:
```
The following submission will be tested:

# Bitte schreiben Sie hier Ihren Programmcode.

# Teil a) Alphabet entwirren: Position jedes Buchstabens aus `alphabet` im sortierten Alphabet
sortiertes_alphabet = sorted(alphabet)
indizes = [sortiertes_alphabet.index(buchstabe) for buchstabe in alphabet]

# Teil b) Häufigkeit zählen
zeichen_anzahl = {buchstabe: wort.count(buchstabe) for buchstabe in alphabet}

sortiertes_alphabet = sorted(alphabet)
indizes = [sortiertes_alphabet.index(buchstabe) for buchstabe in alphabet]

zeichen_anzahl = {buchstabe: wort.count(buchstabe) for buchstabe in alphabet}
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m
```

```python

```