# Addition geschachtelter Listen

Diese Aufgabe wiederholt eine Problemstellung vorheriger Aufgaben (Matrizenaddition) und verallgemeinert sie. Sie können die vorherigen Lösungen als Grundlage verwenden, müssen aber noch eigene Arbeit leisten.

## Problemstellung

In vergangenen Aufgaben haben Sie zwei Listen-von-Listen addiert. Diese sind nicht zwingend rechteckig, haben aber beide die gleiche Form. Das Ergebnis war eine geschachtelte Liste derselben Form.

Nun möchten Sie Ihre Funktion verallgemeinern und *beliebig viele* Listen-von-Listen auf einmal addieren. Diese müssen ebenfalls alle die gleiche Form haben.

## Ihre Aufgabe

Schreiben Sie eine Funktion `add`, die beliebig viele Listen von Listen addiert und eine Liste von Listen zurückgibt, deren Einträge die Summen der jeweiligen Einträge der gegebenen Listen sind.

Als Hilfestellung können Sie sich die Lösungen der anderen Matrizenadditionsaufgaben anschauen. Achten Sie weiterhin darauf, Ihre Schleifen möglichst frei von Zählindizes zu halten, falls Sie Schleifen verwenden müssen. List Comprehensions und eingebaute Funktionen wie `zip`, `sum`, `len`, `any`, `all`, `min`, und `max` können nützlich sein. Lösen Sie diese Aufgabe ohne Verwendung externer Pakete oder Module.

Die bisherigen Anforderungen sollten Sie weiterhin erfüllen:

0. Die Funktion sollte in der Lage sein, zwei rechteckige Matrizen zu addieren.
1. Die Funktion sollte zwei beliebige Listen-von-Listen addieren können, solange beide Listen die gleiche Form haben.
2. Die Funktion soll `None` zurückgeben, falls die gegebenen Listen von Listen nicht die gleiche Form haben.

Aber auch die folgenden Anforderungen sollen erfüllt sein:

3. Die Funktion soll eine beliebige Zahl von Listen-von-Listen addieren können. Diese werden allesamt als positionelle Argumente an die Funktion übergeben.
4. Stellen Sie sicher, dass Ihre Funktion damit umgehen kann, wenn dasselbe veränderliche Objekt mehrfach auftaucht.

<!-- grade_id: cell-fffc2826f313b01a, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.
```

<!-- grade_id: cell-def19781878f4647, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(add)
cases = pathlib.Path('tests') / 'Matrizenaddition'
args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']

assert pytest.main([*args, cases / 'rechteckige_listen.yml']) is pytest.ExitCode.OK
assert pytest.main([*args, cases / 'beliebige_listen.yml']) is pytest.ExitCode.OK
assert pytest.main([*args, cases / 'beliebige_zahl_an_listen.yml']) is pytest.ExitCode.OK
assert pytest.main([*args, cases / 'pruefe_auf_gleiche_form.yml']) is pytest.ExitCode.OK
assert pytest.main([*args, cases / 'mutable_elemente.yml']) is pytest.ExitCode.OK
```