# Das Skalarprodukt

Das Skalarprodukt ist eine grundlegende Operation in der linearen Algebra. Es nimmt zwei Vektoren und gibt einen Skalar zurück. In diesem Tutorial werden Sie lernen, wie Sie das Skalarprodukt in Python berechnen können.



## Grundlage

Das Skalarprodukt wird häufig durch $\langle\cdot,\cdot\rangle$ oder durch das Multiplikationssymbol "$\cdot$" angegeben. Für zwei Vektoren **a** und **b** also beispielsweise durch $\langle\mathbf{a},\mathbf{b}\rangle$ oder $\mathbf{a}\cdot\mathbf{b}$.  Die Standardform des Skalarprodukts kann mit der folgenden Formel berechnet werden:
$$
\mathbf{a} \cdot \mathbf{b} = \sum_{i=1}^n a_i \cdot b_i,
$$
wobei $a_i$ und $b_i$ die $i$-ten Elemente der Vektoren $\mathbf{a}$ und $\mathbf{b}$ sind und $n$ die Länge beider Vektoren wieder gibt.

## Tutorial

### Einfache Implementierung

In Python können Sie das Skalarprodukt auf verschiedene Weisen berechnen, 
welche im Folgendem kurz vorgestellt werden. Wir werden hierbei von der "unpythonisch"sten Variante zu der "pythoischsten" Variante, welche Sie mit Ihren momentanen Kenntnissen implementieren können, kommen. 

Wir werden die Vektoren in Form von Listen betrachten, beispielsweise

```python
a = [1, 2, 3]
b = [4, 5, 6]
```

#### Methode 1: indexbasierte `for`-Schleife

Sie können das Skalarprodukt berechnen, indem Sie über die Elemente der Listen iterieren und die Produkte summieren. Da Sie über zwei Vektoren iterieren, liegt der Ansatz einer indexbasierten `for`-Schleife nahe:

```python
ergebnis = 0
for i in range(len(a)):
    ergebnis += a[i] * b[i]

ergebnis
```

Output:
```
32
```

Dieser Ansatz hat neben einer schlechten Lesbarkeit noch Probleme mit dem Grenzfall `len(a)`$>$ `len(b)`, welcher in einem `IndexError` resultiert. 
(Streng genommen gibt es rein mathematisch gesehen aber auch kein Skalarprodukt für unterschiedliche Dimensionen).

#### Methode 2: `zip` in einer `for`-Schleife


Eine weitere Möglichkeit, das Skalarprodukt zu berechnen, besteht darin, die `zip`-Funktion in einer `for`-Schleife zu verwenden. Hierbei iterieren
Sie direkt über die beiden Listen $a$ und $b$. Dieser Ansatz hat eine höhere Lesbarkeit, außerdem wird im Fall unterschiedlich langer Listen automatisch nur bis zu der kürzeren Länge iteriert. 

```python
ergebnis = 0

for ai, bi in zip(a, b):
    ergebnis += ai * bi

ergebnis
```

Output:
```
32
```

### Fortgeschrittene Methoden


#### Mit einer List Comprehension und der `sum`-Funktion

Eine weitere Möglichkeit, das Skalarprodukt zu berechnen, besteht darin, eine Comprehension zu verwenden. Falls Sie diese noch nicht kennen, werden sie Sie in Kürze in der Vorlesung kennenlernen. 

Dieser Ansatz verwendet die `sum`-Funktion. Mittels einer Comprehension werden die einzelnen Produkte innerhalb der `sum`-Funktion berechnet. Hierdurch erhalten Sie das Ergebnis ohne irgendwelche *Nebeneffekte*, da weder die Variablen `ai` und `bi` noch ein Container mit den Produkten nach der Ausführung existieren. Außerdem zeichnet sich dieser Ansatz durch eine hohe Lesbarkeit aus und ist wahrscheinlich der *pythonischste* für die Berechnung das Skalarprodukts.


```python
sum(ai * bi for ai, bi in zip(a, b))
```

Output:
```
32
```

Zusätzlich könnte man mit `zip(a, b, strict=True)` erzwingen, dass `a` und `b` gleich lang sind und sonst einen Fehler erzeugen.

#### Mit der `map`-Funktion und einer `lambda`-Funktion


Eine weitere Möglichkeit wäre, die `map`-Funktion zu verwenden. Da es allerdings keine eingebaute Produkt-Funktion gibt, müssen wir hierfür eine `lambda`-Funktion verwenden. Dieser Ansatz hat allerdings keine ersichtlichen Vorteile gegenüber der vorherigen Methode und eine schlechtere Lesbarkeit.

```python
sum(map(lambda ai, bi: ai * bi, a, b))
```

Output:
```
32
```

### Mit Hilfe von außen

#### Produkt-Funktion aus Standardlibrary-Modul `operator`

In der Standardlibrary (Standard-Bibliothek der Sprache Python) existiert das Modul `operator`. Dieses enthält die Multiplikation als Operator:

```python
import operator as op

sum(map(op.mul, a, b))
```

Output:
```
32
```

#### Produkt-Funktion aus Standardlibrary-Modul `math`

Im Modul `math` der Standardlibrary existiert die Funktion `prod`, die das Produkt beliebig vieler Werte berechnet:

```python
import math

sum(map(math.prod, zip(a, b)))
```

Output:
```
32
```

Mit Python 3.12 wurde die Funktion `math.sumprod` hinzugefügt, die das Problem sofort löst:

```python
math.sumprod(a, b)
```

Output:
```
32
```

#### Externe Pakete: Numpy, Pandas, etc.

In der Praxis würde man für alle außer die einfachsten Anwendungszwecke ein Paket für lineare Algebra oder Datenanalyse verwenden. Wir implementieren das Skalarprodukt mit Python-`built-ins` rein zu Übungszwecken. `numpy`-Arrays beispielsweise haben eine Methode `dot` und überladen den Operator `@` zum Berechnen von Skalarprodukten.

## Ihre Aufgabe

Implementieren Sie nun das Skalarprodukt für zwei *Listen* oder *Tuple* `a` und `b`. Im Fall von unterschiedlich langen Listen soll die Länge der kürzeren verwendet werden. 

```python
a = [1, 2, 3]
b = (4, 5, 6)
```

<!-- grade_id: cell-8feabe3a52e9bb07, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.

ergebnis = sum(ai * bi for ai, bi in zip(a, b))


```

<!-- grade_id: cell-0b5f247f6adb9a0f, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(_i)

args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']
subtasks = ['integer_tuples','float_tuples','arbitrary_iterables']

assert all(
    pytest.main([*args, pathlib.Path('tests') / 'Skalarprodukt' / f'{subtask}.yml']) is pytest.ExitCode.OK
    for subtask in subtasks
)
```

Output:
```
The following submission will be tested:

# Bitte schreiben Sie hier Ihren Programmcode.

ergebnis = sum(ai * bi for ai, bi in zip(a, b))
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m        [100%][0m
[33ms[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[0m[32m.[
```