# Variable Capturing

Diese Aufgabe beschäftigt sich mit dem sogenannten *Variable Capturing* in Lambda-Funktionen. Das vorliegende Problem wird hierbei vorerst besprochen, anschließend wird Ihnen eine leichte Übungsaufgabe gestellt. Sie sollten sich hierfür bereits mit Lambda-Funktionen beschäftigt haben.

## Problem

Sie möchten eine Lambda-Funktion definieren, welche neben einer übergebenen Variable noch andere, zuvor definierte Variablen verwendet. Den Wert dieser Variable möchten Sie *einfangen*.

Schauen Sie sich folgendes Beispiel an:

```python
x = 10
a = lambda y: x +y
x = 20
b = lambda y: x + y
```

Was würden Sie bei der Auswertung von `a(10)` und `b(10)` für ein Ergebnis erwarten? 

```python
a(10), b(10)
```

Wie Sie, vielleicht unerwartet, sehen können, hat sich die Funktion `a` den Wert von `x` nicht gemerkt. Das liegt daran, dass `x` eine freie Variable ist, dessen Wert bei der **Ausführung** der Lambda-Funktion zugeordnet wird. Der Wert von `x` in der Lambda-Funktion wird also immer der sein, der `x` gerade zugeordnet ist. 

Einige Beispiele:

```python
x = 15
a(10)
```

```python
x = 12.5
a(10)
```

```python
del x
a(10)
```

```python
x = "abc"
a(10)
```

```python
x = 5 
a(10)
```

## Lösung

Soll Wert von `x` bei der Definition von `a` wirklich *gefangen* werden, so können Sie diesen als optionalen Parameter mit Standardwert definieren:

```python
x = 10
a = lambda y, x=x: x+y
```

```python
x = 15
a(10)
```

Nun ist der Wert zum Zeitpunkt der Ausführung egal, innerhalb von `a` hat `x` standardmäßig den Wert 10. Dieser kann wie bei jeder Funktion mit optionalen Parametern beim Funktionaufruf geändert werden:

```python
a(10, 15), a(10, x = 15), a(10, x=x)
```

Problemstellungen wie diese können beispielsweise auftreten, wenn Sie Lambda-Funktionen in einer Schleife erstellen, welche auf der Schleifenvariable basieren. 

## Ihre Aufgabe

Sie möchten eine Liste von Lambda-Funktionen erstellen, welche eine Zahl `x` mit einem immer größer werdenden Wert `a` addiert. Dafür verwenden Sie eine List-Comprehension:

```python
funcs = [lambda x: x + a for a in range(5)]
for f in funcs:
    print(f(0))
```

Die Funktionen erinnern sich nicht an den Wert von `a`. Lösen Sie dieses Problem. 

Verpacken Sie Ihre Lösung in eine Funktion namens `create_list_of_functions`, welche einen Parameter `n` bekommt und die Liste von Funktionen zurückgibt. Diese soll bis `range(n)` statt bis `range(5)` erstellt werden.

<!-- grade_id: cell-bc95e672c3e1ef13, points: 0, grade: false, solution: true -->
```python
# Bitte schreiben Sie hier Ihren Programmcode.
```

<!-- grade_id: cell-497ebdbcbba60c23, points: 1, grade: true, solution: false -->
```python
import pathlib

import pytest
from pytest_nbgrader import loader

loader.Submission.submit(create_list_of_functions)
cases = pathlib.Path('tests') / 'Variable_Capturing'
test_resource = cases / 'tests.py::TestVariableCapturing::test_'
args = ['-qq', '-x', '-W', 'ignore::_pytest.warning_types.PytestAssertRewriteWarning', '--cases']

assert pytest.main([*args, cases / 'variable_capturing.yml', f'{test_resource}variable_capturing']) is pytest.ExitCode.OK
```